# Assignment 2

More on Accumulation

Write code that **uses `map`** to assign to the variable `map_testing` all the elements from `lst_check`, prepending the string `"Fruit: "` to the beginning of each element. *Note: be sure to use `list` to convert the result from `map` into a list.*

<!-- grade_id: cell-33f705472409f897, points: 0, grade: false, solution: true -->
```python
lst_check = ['plums', 'watermelon', 'kiwi', 'strawberries', 'blueberries', 'peaches', 'apples', 'mangos', 'papaya']

map_testing = list(map(lambda s: "Fruit: " + s, lst_check))

assert 'Fruit: plums' in map_testing
assert 'Fruit: kiwi' in map_testing

```

<!-- grade_id: cell-9270b9c7ad936d05, points: 1, grade: true, solution: false -->
```python
assert 'Fruit: plums' in map_testing
assert 'Fruit: kiwi' in map_testing

```

Below, we have provided a list of strings called `countries`. **Use `filter`** to produce a list called `b_countries` that only contains the strings from countries that begin with `'B'`. *Note: be sure to use `list` to convert the result from `filter` into a list.*

<!-- grade_id: cell-728dbbed11e70922, points: 0, grade: false, solution: true -->
```python
countries = ['Canada', 'Mexico', 'Brazil', 'Chile', 'Denmark', 'Botswana', 'Spain', 'Britain', 'Portugal', 'Russia', 'Thailand', 'Bangladesh', 'Nigeria', 'Argentina', 'Belarus', 'Laos', 'Australia', 'Panama', 'Egypt', 'Morocco', 'Switzerland', 'Belgium']

# Use filter with a lambda function to check if a country starts with 'B'
# Then, convert the filter object to a list
b_countries = list(filter(lambda country: country.startswith('B'), countries))

assert 'Brazil' in b_countries
assert 'Canada' not in b_countries

```

<!-- grade_id: cell-a3a6ce379963c355, points: 1, grade: true, solution: false -->
```python
assert 'Brazil' in b_countries
assert 'Canada' not in b_countries

```

Below, we have provided a list of tuples that contain the names of Game of Thrones characters. The first item in each tuple is their last (family) name. The second item is their first (given) name. Using a **list comprehension**, create a list of strings called `first_names` that contains only the first names of everyone in the original list.

<!-- grade_id: cell-6d4a02615ad9dc17, points: 0, grade: false, solution: true -->
```python
people = [('Snow', 'Jon'), ('Lannister', 'Cersei'), ('Stark', 'Arya'), ('Stark', 'Robb'), ('Lannister', 'Jamie'), ('Targaryen', 'Daenerys'), ('Stark', 'Sansa'), ('Tyrell', 'Margaery'), ('Stark', 'Eddard'), ('Lannister', 'Tyrion'), ('Baratheon', 'Joffrey'), ('Bolton', 'Ramsey'), ('Baelish', 'Peter')]

# Use a list comprehension to extract the first name (second element) from each tuple
first_names = [person[1] for person in people]

assert 'Jon'  in first_names
assert 'Snow' not in first_names

```

<!-- grade_id: cell-1d94c540a0067985, points: 1, grade: true, solution: false -->
```python
assert 'Jon'  in first_names
assert 'Snow' not in first_names

```

Use a **list comprehension** to create a list called `lst2` that "doubles" each element in the list, `lst`. Use `* 2` to "double" numbers (`5 * 2 == 10`), strings (`"xyz" * 2 == "xyzxyz"`), and lists (`['a', 'b'] * 2 == ['a', 'b', 'a', 'b']`).

<!-- grade_id: cell-3d3befbc2c890e0a, points: 0, grade: false, solution: true -->
```python
lst = [["hi", "bye"], "hello", "goodbye", [9, 2], 4]

# Use a list comprehension to "double" each element
lst2 = [item * 2 for item in lst]

assert "hellohello" in lst2

```

<!-- grade_id: cell-f1ef42d126ceaad3, points: 1, grade: true, solution: false -->
```python
assert "hellohello" in lst2

```

Below, we have provided a list of tuples that contain students' names and their final grades in PYTHON 101. Using a **list comprehension**, create a new list `passed` that contains the *names* of students who passed the class (had a final grade of `70` or greater). `passed` should be a list of strings (names without final grades) in the same order as `students`.

<!-- grade_id: cell-0026cd972495389b, points: 0, grade: false, solution: true -->
```python
students = [('Tamara', 95), ('Linda', 63), ('Carl', 70), ('Bob', 100), ('Raymond', 50), ('Sue', 75)]

# Use a list comprehension to filter students who passed (grade >= 70)
# and then extract only their names
passed = [name for name, grade in students if grade >= 70]

assert 'Tamara' in passed
assert 'Linda' not in passed

```

<!-- grade_id: cell-7365646af3677b07, points: 1, grade: true, solution: false -->
```python
assert 'Tamara' in passed
assert 'Linda' not in passed

```

Write code using `zip` and `filter` so that these lists (`l1` and `l2`) are combined into one big list of pairs and assigned to the variable `opposites` if both elements of the pair are longer than `3` characters. *Note: be sure to use `list` to convert the result into a list.*

<!-- grade_id: cell-4a5785241c831a5a, points: 0, grade: false, solution: true -->
```python
l1 = ['left', 'up', 'front']
l2 = ['right', 'down', 'back']

# 1. Zip the lists to create pairs: [('left', 'right'), ('up', 'down'), ('front', 'back')]
# 2. Filter these pairs:
#    The lambda function checks if the length of the first element (x[0]) is > 3
#    AND if the length of the second element (x[1]) is > 3.
# 3. Convert the filter object to a list.
opposites = list(filter(lambda x: len(x[0]) > 3 and len(x[1]) > 3, zip(l1, l2)))

assert ('left', 'right') in opposites
# omit if either word is shorter than 3 characters...
assert ('up', 'down') not in opposites

```

<!-- grade_id: cell-077fa8cd007e9e9c, points: 1, grade: true, solution: false -->
```python
assert ('left', 'right') in opposites
# omit if either word is shorter than 3 characters...
assert ('up', 'down') not in opposites

```

Below, we have provided a `species` list and a `population` list. **Use `zip`** to combine these lists into one list of tuples called `pop_info`. From this list, create a new list called `endangered` that contains the names of species whose populations are below `2500`.

<!-- grade_id: cell-5fcdbdf62ecae78e, points: 0, grade: false, solution: true -->
```python
species = ['golden retriever', 'white tailed deer', 'black rhino', 'brown squirrel', 'field mouse', 'orangutan', 'sumatran elephant', 'rainbow trout', 'black bear', 'blue whale', 'water moccasin', 'giant panda', 'green turtle', 'blue jay', 'japanese beetle']

population = [10000, 90000, 1000, 2000000, 500000, 500, 1200, 8000, 12000, 2300, 7500, 100, 1800, 9500, 125000]

# Combine species and population into a list of tuples using zip
pop_info = list(zip(species, population))

# Create a list of endangered species using a list comprehension
# Iterate through pop_info, unpack each tuple into 's' (species) and 'p' (population)
# Include 's' in the new list if 'p' (population) is less than 2500
endangered = [s for s, p in pop_info if p < 2500]


assert ('golden retriever', 10000) in pop_info
assert 'golden retriever' not in endangered
assert 'black rhino' in endangered

```

<!-- grade_id: cell-78396308d9912e18, points: 1, grade: true, solution: false -->
```python
assert ('golden retriever', 10000) in pop_info
assert 'golden retriever' not in endangered
assert 'black rhino' in endangered

```